# MSCF 46982 Market Microstructure and Algorithmic Trading
# Fall 2025 Mini 2

Before you turn this assignment in, make sure everything runs as
expected. First, **restart the kernel**: by selecting
Kernel$\rightarrow$Restart Kernel from the menubar (or the `0,0`
keyboard shortcut) and then **run all cells** by selecting
Run$\rightarrow$Run All Cells from the menubar.

Make sure you replace all instances `YOUR CODE HERE` or "YOUR ANSWER
HERE" with your solution and remove the `notimplemented` exception.


As indicated in the syllabus, you are encouraged to discuss the
material presented in class but not of the specifics of assignments
(including this one).


---

```python
import os
os.environ['PYKX_JUPYTERQ'] = 'true'
os.environ['PYKX_4_1_ENABLED'] = 'true'
import pykx as kx
```

Output:
```
<frozen importlib._bootstrap>:488: PyKXWarning: Unable to connect user QHOME to PyKX QHOME via symlinks.
To permanently disable attempts to create symlinks you can
	1. Set the environment variable "PYKX_IGNORE_QHOME" = True.
	2. Update the file ".pykx-config" using kx.util.add_to_config({'PYKX_IGNORE_QHOME': 'True'})
Error: [WinError 1314] A required privilege is not held by the client: 'C:\\Users\\jacob\\Desktop\\Fin Computing I\\Mkt Micro and Algo\\hw0.ipynb' -> 'C:\\Users\\jacob\\Desktop\\Fin Computing I\\env\\Lib\\site-packages\\pykx\\lib\\4-1-libs\\hw0.ipynb'

PyKX now running in 'jupyter_qfirst' mode. All cells by default will be run as q code. 
Include '%%py' at the beginning of each cell to run as python code.
```

## Anaconda w/ PyKX Installation

To confirm you have installed Anaconda with the kdb+ kernel installed
properly, compute the `md5` and `sha1` hash of 4.1.


### Part A (5 Points)

The `.z.K` internal variable contains the version of kdb+. Evaluate
the below cell, to confirm you are running version 4.1 and display the
md5 and sha1 hashed values of this version. A hash is a mathematical
function that converts an arbitrary-length input to a fixed-length
output.  Hashes are used heavily to prove that two values are
identical while not revealing anything about the original values.
Small changes to the input result in completely different hashed
values.  The longer the hash, the more secure it is.  Bitcoin uses the
SHA-2 (aka SHA-256) algorithm, which creates 256-bit hashes (instead
of the shorter SHA-1 algorithm which uses 160-bit hashes).

The code below stores the md5 and sha1 hashes of the value 3.6. Using
the kdb+ [reference card](https://code.kx.com/q/ref/) and the .Q
[*unreference* card](https://code.kx.com/q/ref/dotq/) of internal
functions, update the declaration of `md5hash` and `sha1hash` to have
the correct hashed value of **.z.K** converted to a string.


<!-- grade_id: dotzdotk, points: 0, grade: false, solution: true -->
```python
md5hash:md5 string .z.K
sha1hash: .Q.sha1 string .z.K
```

<!-- grade_id: correct_dotzdotk, points: 5, grade: true, solution: false -->
```python
assert:{if[not x~y;'"expecting '",(-3!x),"' but found '",(-3!y),"'"]}
assert[4.1] .z.K
assert[0xefd70e52d5b15f87444582f6d779bdfa] -15! raze string md5hash
assert[0x35193e5e3999b0539c96695131c0ff2f76dda7de] -33! raze string sha1hash
```