# MSCF 46982 Market Microstructure and Algorithmic Trading
# Fall 2025 Mini 2

Before you turn this assignment in, make sure everything runs as
expected. First, **restart the kernel**: by selecting
Kernel$\rightarrow$Restart Kernel from the menubar (or the `0,0`
keyboard shortcut) and then **run all cells** by selecting
Run$\rightarrow$Run All Cells from the menubar.

Make sure you replace all instances `YOUR CODE HERE` or "YOUR ANSWER
HERE" with your solution and remove the `notimplemented` exception.


As indicated in the syllabus, you are encouraged to discuss the
material presented in class but not of the specifics of assignments
(including this one).


---

```python
import os
os.environ['PYKX_JUPYTERQ'] = 'true'
os.environ['PYKX_4_1_ENABLED'] = 'true'
import pykx as kx

```

Output:
```
<frozen importlib._bootstrap>:488: PyKXWarning: Unable to connect user QHOME to PyKX QHOME via symlinks.
To permanently disable attempts to create symlinks you can
	1. Set the environment variable "PYKX_IGNORE_QHOME" = True.
	2. Update the file ".pykx-config" using kx.util.add_to_config({'PYKX_IGNORE_QHOME': 'True'})
Error: [WinError 1314] A required privilege is not held by the client: 'C:\\Users\\jacob\\Desktop\\Fin Computing I\\Mkt Micro and Algo\\hw0.ipynb' -> 'C:\\Users\\jacob\\Desktop\\Fin Computing I\\env\\Lib\\site-packages\\pykx\\lib\\4-1-libs\\hw0.ipynb'

PyKX now running in 'jupyter_qfirst' mode. All cells by default will be run as q code. 
Include '%%py' at the beginning of each cell to run as python code.
```

## Fibonacci Sequence
While kdb+ comes with a `while` operator, it is rarely used and
frowned upon as good form.  This problem takes you through a popular
example of how iteration is performed in q using adverbs (now called
iterators). An introduction to kdb+ and an example of the Fibonacci
sequence can also be viewed in [Q for All][Q4A]

[Q4A]: https://code.kx.com/v2/learn/q-for-all/ "Q For All"


### Part A (3 Points)

Write a function that returns the sum of the last 2 elements of a list.

Hint: use the `#` operator.

<!-- grade_id: sumlasttwo, points: 0, grade: false, solution: true -->
```python
/ returns the sum of the last two elements
sumlasttwo:{ sum -2#x }

```

Your function should return `8` for `x:0 1 1 2 3 5`. Check that it does:

```python
sumlasttwo 0 1 1 2 3 5
```

Output:
```
8
```

<!-- grade_id: correct_sumlasttwo, points: 1, grade: true, solution: false -->
```python
assert:{if[not x~y;'"expecting '",(-3!x),"' but found '",(-3!y),"'"]}
/ check that sumlasttwo produces the correct output for several inputs
assert[2] sumlasttwo 0 1 1
assert[3] sumlasttwo 0 1 1 2
assert[5] sumlasttwo 0 1 1 2 3
assert[8] sumlasttwo 0 1 1 2 3 5
assert[13] sumlasttwo 0 1 1 2 3 5 8
```

<!-- grade_id: correct_sumlasttwo_ec, points: 1, grade: true, solution: false -->
```python
/ check that the function is implemented efficiently
assert[1b] 131376>=0N!last system "ts sumlasttwo til 10000"
```

### Part B (2 Points)

Write a function `fib` that appends the sum of the last two elements
of a list to the original list.

Hint: use compound assignment

<!-- grade_id: fib, points: 0, grade: false, solution: true -->
```python
/ returns the sum of the last two elements
fib:{x,: sumlasttwo x; x}
```

Your function should return `0 1 1 2 3 5 8` for `x:0 1 1 2 3 5`.
Check that it does:


```python
fib 0 1 1 2 3 5
```

Output:
```
0 1 1 2 3 5 8
```

<!-- grade_id: correct_fib, points: 2, grade: true, solution: false -->
```python
/ check that fib produces the correct output for several inputs
assert[0 1 1] fib 0 1
assert[0 1 1 2] fib 0 1 1
assert[0 1 1 2 3] fib 0 1 1 2
assert[0 1 1 2 3 5] fib 0 1 1 2 3
assert[0 1 1 2 3 5 8] fib 0 1 1 2 3 5
assert[0 1 1 2 3 5 8 13] fib 0 1 1 2 3 5 8
```

### Part C (3 Points)

Using your `fib` function and the `/` [iterator][iterators], write a
new function `fibonacci` which returns the first n elements of the
Fibonacci sequence after `0 1`.

[iterators]: https://code.kx.com/v2/ref/iterators/ "Q Iterators"


<!-- grade_id: fibonacci, points: 0, grade: false, solution: true -->
```python
/ returns the sum of the last two elements
fibonacci:{[n] fib/[n; 0 1]}

```

Your function should return `0 1 1 2 3 5` for `n:4`.  Check that it does:

```python
fibonacci 4
```

Output:
```
0 1 1 2 3 5
```

<!-- grade_id: correct_fibonacci, points: 2, grade: true, solution: false -->
```python
/ check that fibonacci produces the correct output for several inputs
assert[0 1] fibonacci 0
assert[0 1 1] fibonacci 1
assert[0 1 1 2] fibonacci 2
assert[0 1 1 2 3] fibonacci 3
assert[0 1 1 2 3 5] fibonacci 4
assert[0 1 1 2 3 5 8] fibonacci 5
assert[1298777728820984005] last fibonacci 100
```

<!-- grade_id: correct_fibonacci_ec, points: 1, grade: true, solution: false -->
```python
/ check that your function is implemented efficiently
assert[1b] 2368>=0N!last system"ts fibonacci 100"
```

## Sharpe Ratio (3 Points)

A common performance statistic for investors is the **Sharpe ratio**.
It is defined as the average profit (above the risk free rate) divided
by the risk (or standard deviation) of achieving that profit.

$$\frac{\text{average daily return} - \text{daily risk free
return}}{\text{daily standard deviation of returns}}$$

Many long/short strategies are self financing and the cost of capital
is therefore ignored.  And instead of computing the **Sharpe ratio**
using returns, it is common to use the actual profits (aka: pnl).

$$\frac{\text{average daily pnl}}{\text{daily standard deviation of pnl}}$$

Using the `avg` and `sdev` operator write a `sharpe` function that
computes the **Sharpe ratio** of a given series of daily pnl results.


Hint: make sure to compute the ratio of the average and the standard
deviation -- not the average of the ratios.


<!-- grade_id: sharpe, points: 0, grade: false, solution: true -->
```python
/ compute the Sharpe ratio of a list of values
sharpe:{(avg x) % sdev x}

```

Given a series of daily profits (and losses) `20 -10 30 10 15 -5 20 50
30 20`, your function should return `1.037311`.  Check that it does:


```python
sharpe 20 -10 30 10 15 -5 20 50 30 20
```

Output:
```
1.037311
```

<!-- grade_id: correct_sharpe, points: 2, grade: true, solution: false -->
```python
/ check that sharpe produces the correct results
assert[0.5316765e] "e"$sharpe -4 -1 -2 -3 5 4 3 10 5 10
assert[0.4685095e] "e"$sharpe -100 100 -300 50 400 -50 1000 50 500
```

<!-- grade_id: correct_sharpe_ec, points: 1, grade: true, solution: false -->
```python
/ check that your function is implemented efficiently
assert[1b] 147920>=0N!last system "ts sharpe 10000?1f"
```

## Drawdown (3 Points)

Another common performance statistic for investors is the
**drawdown**.  It is defined as the loss from peak to trough.  If a
stock starts at 100, rises to 120 and falls to 90, the drawdown would
be 30 (120-90) or 25% (30/120).  If the price rises again to 110, the
drawdown would be 8.33% (10/120).  And finally, if the price moved
beyond 120 to 130, the draw down would be 0.


Using the `maxs` and `%` operators, write a `drawdown` function that
computes a time-series of the drawdown statistic.


<!-- grade_id: drawdown, points: 0, grade: false, solution: true -->
```python
/ compute the drawdown (as a percentage) of a list of values
drawdown:{ ((maxs x) - x) % 1f* maxs x }
```

Given a series of prices `100 101 100 99 98 100 102 104 103`, your
function should return `0 0 0.00990099 0.01980198 0.02970297
0.00990099 0 0 0.009615385`.  Check that it does:


```python
drawdown 100 101 100 99 98 100 102 104 103
```

Output:
```
0 0 0.00990099 0.01980198 0.02970297 0.00990099 0 0 0.009615385
```

<!-- grade_id: correct_drawdowns, points: 2, grade: true, solution: false -->
```python
/ check that drawdown produces the correct results
assert[0 0 0.5 0.25 0 0.15] drawdown 100 120 60 90 120 102
assert[0 0.1 0.2 0.3 0.4 0.5 0.1 0.2 0] drawdown 100 90 80 70 60 50 90 80 100
```

<!-- grade_id: correct_drawdowns_ec, points: 1, grade: true, solution: false -->
```python
/ check that your function is implemented efficiently
assert[1b] 5456>=0N!last system "ts drawdown 100?1000"
```

### Relative Strength Index (3 Points)

The Relative Strength Index (RSI) is a technical momentum indicator
bounded between 0 and 100 and used to detect reversion opportunities.

It is computed as:
$$
\text{RSI} = 100-\frac{100}{1+\text{RS}}
$$

where RS is defined as the ratio of average up moves divide by the
average down moves.  "Average move" can be defined in many ways:
Exponential Moving Average, Smoothed Moving Average or even a Simple
Moving Average.

In this example we will use the Simple Moving Average:
$$
\text{RS} = \frac{\frac{1}{n}\sum_{t=1-n}^0\max(0,P_t-P_{t-1})}{\frac{1}{n}\sum_{t=1-n}^0\max(0,P_{t-1}-P_t)}
$$

Using the `mavg`, `prev` and `|` operators, write an `rsi` function
that computes the Relative Strength Index


<!-- grade_id: rsi, points: 0, grade: false, solution: true -->
```python
/ compute the relative strength index over (w)indow of a list of numbers (x)
rsi:{[w;x] 
     ups:w mavg 0|x-prev x; 
     downs:w mavg 0|(prev x)-x; 
     rs:ups%downs; 
     100-100%1+rs}
```

Given a window size of 10 and the same series of prices `100 101 100
99 98 100 102 104 103`, your function should return `0n 100 50
33.33333 25 50 62.5 70 63.63636`.  Check that it does:


```python
rsi[10] 100 101 100 99 98 100 102 104 103
```

Output:
```
0n 100 50 33.33333 25 50 62.5 70 63.63636
```

<!-- grade_id: correct_rsi, points: 2, grade: true, solution: false -->
```python
/ check that mdd produces the correct value for a few inputs
assert[0n 100 100 100 100 100 100 100 100 100 100] rsi[10] 1 2 3 4 5 6 7 8 9 10 11
assert[0n 0 0 0 0 0 0 0 0 0 0] rsi[10] neg 1 2 3 4 5 6 7 8 9 10 11
assert[0n 0 75 75 80 40 40 12.5 10] rsi[5] 100 99 102 102 103 98 97 96 94
assert[0n 0 75 75 80 40 25 40 25] rsi[10] 100 99 102 102 103 98 92 96 84
```

<!-- grade_id: correct_rsi_ec, points: 1, grade: true, solution: false -->
```python
/ check that your function is implemented efficiently
assert[1b] 1552>=0N!last system"ts rsi[10] 100 99 102 102 103 98 92 96 84"
```

## Dictionaries, Tables and Keyed Tables

Kdb+ is a vector programming language, and as such, vectors are
everywhere.  "Atomic" functions work equally well on atoms and
vectors.  Remember:
- A dictionary is a pair of vectors.
- A table can be treated as a list of dictionaries (though it is
  stored in memory as a flipped dictionary of lists).
- A Keyed table is a pair of tables and can be thought of as a
  dictionary mapping one table to another.

The following questions exercise your knowledge of dictionaries and tables.

### Part A (2 Points)

Using the `!` operator, create a dictionary mapping each lowercase
letter of the alphabet to the numbers 1 to 26 and store it in the
variable `d`.


Hint: Use variables in the .Q namespace to make your life easier

<!-- grade_id: dictionary, points: 0, grade: false, solution: true -->
```python
d:(.Q.a)!(1+til 26)
```

<!-- grade_id: correct_dictionary, points: 2, grade: true, solution: false -->
```python
/ check there are 26 elements in the dictionary
assert[26] count d
/ check the key is a string vector
assert[10h] type key d
/ check the values sum to 251
assert[351] sum d
/ check that carnegie mellon university sums to 295
assert[295] sum d "carnegie mellon university"

```

### Part B (2 Points)

Using `([]a;b;c)` notation, generate a table `t` with three columns:
`date`, `sym` and `price`.  The table should have dates ranging from
2021.01.01 thru 2021.12.31, the `sym` column should all have the value
`BAC`, and the `price` column should a list of random floating point
numbers.


<!-- grade_id: table, points: 0, grade: false, solution: true -->
```python
t:([]date:(2021.01.01+til 365);sym:(365#`BAC);price:(365?100f))
```

<!-- grade_id: correct_table, points: 2, grade: true, solution: false -->
```python
/ check there are 365 rows in the table
assert[365] count t
/ check the three columns are date,sym,price
assert[`date`sym`price] cols t
/ check the dates range from 2021.01.01 thru 2021.12.31
assert[2021.01.01 2021.12.31] (first;last)@\:t`date
/ check the sym column has a single value: `BAC
assert[1#`BAC] distinct t`sym
/ check that the price column has floats
assert[9h] type t`price
/ check that there are no null values in the price column
assert[1b] all not null t`price
```

### Part C (2 Points)

Using `([a]b;c)` notation, generate a table `t` with three columns:
`andrewid`, `fname`, `lname`.  The `andrewid` column should be of type
symbol and the `fname` and `lname` column should be of type string
(character vector). The table should have `andrewid` as its primary
key. You may put as many rows as you want, but it should have the
following 2 records:

```
([andrewid: jpsaris; fname: "nick"; lname: "psaris"])
([andrewid: cspatt; fname: "chester"; lname: "spatt"])
```


<!-- grade_id: keytable, points: 0, grade: false, solution: true -->
```python
t:([andrewid:`jpsaris`cspatt] fname:("nick";"chester");lname:("psaris";"spatt"))

```

<!-- grade_id: correct_keytable, points: 2, grade: true, solution: false -->
```python
/ check the table column names are andrewid, fname and lname
assert[`andrewid`fname`lname] cols t
/ check the primary key andrewid
assert[1#`andrewid] keys t
/ check the column types are correct
assert["sCC"] value[meta t]`t
/ check the name of jpsaris is nick psaris
assert[("nick";"psaris")] t[`jpsaris;`fname`lname]
/ check the name of dseppi is duane sepppi
assert[("chester";"spatt")] t[`cspatt;`fname`lname]
```

## Eratosthenes Sieve

The [Sieve of Eratosthenes][SOE] is an algorithm which finds all
primes below a supplied limit by iteratively dropping all numbers that
are multiples of other numbers.

[SOE]: https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes "Sieve of Eratosthenes"

Suppose we are looking for the all primes below 10, for example.  We
start with numbers between 2 and 10.

```
2 3 4 5 6 7 8 9 10
```

We then pass this vector through the sieve by first dropping all
factors of 2.

```
2 3 5 7 9
```

We then pass the result through the next layer of the sieve by
dropping all factors of 3.

```
2 3 5 7
```

We have now completed the sieve of Eratosthenes for the first 10
numbers.  Notice how we only needed to try the values (v) 2 and 3.  It
can be shown that we do not need to pass the vector through a sieve
for values (v) higher than the square root of n (which in our case is
3 because the square root of 10 is 3.162278). Notice also, that we did
not have to filter the number 6 in the second pass of the filter, even
though 6 is a multiple of 3.  The first pass had already removed it
and we can optimize the sieve by only checking values (v) between v^2
and n. In other words, for v=2 we begin filtering with 4 and for v=3
we begin filtering with 9.

Feel free to search [Rosetta Code][RC] which has implementations in
other languages.

[RC]: https://rosettacode.org/wiki/Sieve_of_Eratosthenes
"Sieve of Eratosthenes on Rosetta Code"

### Part A (2 Points)

Write a function `sieve` that accepts an integer `n` and returns all
primes less than or equal to that value.

- Begin by generating a vector of boolean values `b` of length `n+1`
  (all set to true) that will be passed through the sieve
- Use the `sqrt` and `floor` operator to generate a vector of values
  `v` that ranges between 2 and the square root of `n` that will be
  iteratively applied through the sieve
- Write a separate (or inner) function `f` that accepts the boolean
  vector `b` and a single value `v`. The function should return the
  boolean vector after applying the sieve for the provided value. In
  other words, it should set all indices that are a multiple of v to
  false (`0b`).
- Using the the `/` adverb (aka iterator), iteratively apply the sieve
  to the boolean vector `b` for all values `v`
- Use the `where` and `_` drop operators to obtain the list of primes
  `p` from the final boolean vector `b`
- HINT: For efficiency, use the `if` operator to only apply the sieve
  against values `v` that are known to be prime
- HINT: Use the integer division operator `div` to help generate the
  list of indices between `v*v` and `n` which should be filtered in
  your function `f`


<!-- grade_id: sieve, points: 0, grade: false, solution: true -->
```python
sieve:{[n]
    b:(n+1)#1b;
    v:2+til -1+floor sqrt n;
    f:{[b;v]
       if[b[v];
          b[(v*v)+(v*til 1+((count[b]-1)-v*v)div v)]:0b
        ];
        b
    };
    b:f/[b;v];
    2_where b}
```

Your function should return `2 3 5 7` for the list of primes less than
or equal to 10.  Check that it does:


```python
sieve 10
```

Output:
```
2 3 5 7
```

<!-- grade_id: correct_sieve, points: 2, grade: true, solution: false -->
```python
/ check primes through 10
assert[2 3 5 7] sieve 10
/ check primes through 100
assert[2 3 5 7 11 13 17 19 23 29 31 37 41 43 47 53 59 61 67 71 73 79 83 89 97] sieve 100
/ check number of primes through 1,000
assert[168] count sieve 1000
/ check last prime through 1,000
assert[997] last sieve 1000
```

<!-- grade_id: correct_sieve_ec, points: 1, grade: true, solution: false -->
```python
/ check that your function is implemented efficiently
assert[1b] 10494576>=0N!last system "ts sieve 1000000"
```

Output:
```
File c:\Users\jacob\Desktop\Fin Computing I\env\Lib\site-packages\pykx\nbextension.py:181, in q(instructions, code)
    179     if save:
    180         print(f"Cell contents not saved to '{path}' due to error during execution/saving.")
--> 181     raise e
    182 finally:
    183     if issubclass(type(_q), kx.QConnection):

File c:\Users\jacob\Desktop\Fin Computing I\env\Lib\site-packages\pykx\nbextension.py:169, in q(instructions, code)
    167     if debug or kx.config.pykx_qdebug:
    168         print(r['trc'])
--> 169     raise kx.QError(r['res'].py().decode())
    170 else:
    171     display(r['res']) if displayRet else print(r['res'])

QError: expecting '1b' but found '0b'
```